May 15 00:13:12 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - load average: 0.24, 2.03, 2.23
May 15 00:08:12 vespa nagios3: SERVICE ALERT: localhost;Current Load;WARNING;HARD;4;WARNING - load average: 3.67, 4.18, 2.72
May 15 00:07:12 alien nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;3;WARNING - load average: 3.69, 4.36, 2.68
May 15 00:06:12 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 333 processes
May 15 00:06:12 cats nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 3.52, 4.44, 2.58
May 15 00:05:12 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 4.30, 4.74, 2.55
May 15 00:04:12 ank-morpork nagios3: SERVICE ALERT: localhost;Current Users;WARNING;HARD;4;USERS WARNING - 25 users currently logged in
May 15 00:03:12 alien nagios3: SERVICE ALERT: localhost;Current Users;WARNING;SOFT;3;USERS WARNING - 25 users currently logged in
May 15 00:02:12 cats nagios3: SERVICE ALERT: localhost;Current Users;WARNING;SOFT;2;USERS WARNING - 25 users currently logged in
May 15 00:01:22 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;CRITICAL;HARD;4;PROCS CRITICAL: 416 processes
May 15 00:01:12 ank-morpork nagios3: SERVICE ALERT: localhost;Current Users;WARNING;SOFT;1;USERS WARNING - 24 users currently logged in
May 16 00:05:36 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 255 processes
May 16 00:04:36 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes
May 16 00:03:36 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes
May 16 00:02:36 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
May 15 23:57:36 cats nagios3: SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 250 processes
May 15 23:47:36 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 254 processes
May 15 23:46:36 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes
May 15 23:45:36 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 258 processes
May 15 23:45:06 cats nagios3: SERVICE ALERT: localhost;Disk Space;OK;HARD;4;DISK OK
May 15 23:44:36 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 306 processes
May 15 23:20:06 alien nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;HARD;4;DISK CRITICAL - /home/croute/ens/etgl/irs/m2/p16/coursrappel/coursrappel is not accessible: Permission denied
May 15 23:19:06 cats nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;3;DISK CRITICAL - /home/croute/ens/etgl/irs/m2/p16/coursrappel/coursrappel is not accessible: Permission denied
May 15 23:18:06 vespa nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;2;DISK CRITICAL - /home/croute/ens/etgl/irs/m2/p16/coursrappel/coursrappel is not accessible: Permission denied
May 15 23:17:06 vespa nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;1;DISK CRITICAL - /home/croute/ens/etgl/irs/m2/p16/coursrappel/coursrappel is not accessible: Permission denied
May 15 21:24:36 cats nagios3: SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 238 processes
May 15 21:21:46 vespa nagios3: SERVICE ALERT: localhost;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
May 15 17:13:54 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;OK;SOFT;4;OK - load average: 1.56, 3.46, 2.53
May 15 17:12:54 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;3;WARNING - load average: 3.23, 4.07, 2.65
May 15 17:11:54 cats nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 6.77, 4.70, 2.74
May 15 17:10:54 vespa nagios3: SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.70, 4.83, 2.64
May 15 09:03:51 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;OK;HARD;4;OK - load average: 0.76, 2.11, 2.29
May 15 08:58:51 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;WARNING;HARD;4;WARNING - load average: 1.79, 4.26, 2.83
May 15 08:57:51 vespa nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;3;WARNING - load average: 3.23, 5.01, 2.96
May 15 08:56:51 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 5.67, 5.76, 3.05
May 15 08:55:51 alien nagios3: SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.22, 6.71, 3.15
May 17 08:51:13 cats nagios3: SERVICE ALERT: localhost;Current Load;OK;SOFT;3;OK - load average: 2.70, 3.64, 1.93
May 17 08:50:13 cats nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 7.24, 4.44, 2.06
May 17 08:49:13 cats nagios3: SERVICE ALERT: localhost;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 16.60, 4.93, 2.04
May 16 16:48:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 256 processes
May 16 16:47:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes
May 16 16:46:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes
May 16 16:45:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes
May 16 16:40:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;4;PROCS OK: 250 processes
May 16 16:39:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes
May 16 16:38:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes
May 16 16:37:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
May 16 16:32:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;4;PROCS OK: 249 processes
May 16 16:31:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes
May 16 16:30:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes
May 16 16:29:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes
May 16 16:14:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 249 processes
May 16 16:13:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes
May 16 16:08:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 249 processes
May 16 15:13:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 253 processes
May 16 15:12:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes
May 16 15:11:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes
May 16 15:10:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes
May 16 14:40:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 250 processes
May 16 14:39:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
May 16 14:29:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;4;PROCS OK: 250 processes
May 16 14:28:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes
May 16 14:27:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes
May 16 14:26:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes
May 16 13:56:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;3;PROCS OK: 250 processes
May 16 13:55:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes
May 16 13:54:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
May 16 13:44:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;3;PROCS OK: 248 processes
May 16 13:43:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 258 processes
May 16 13:42:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes
May 16 13:32:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 250 processes
May 16 13:17:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes
May 16 13:16:54 cats nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes
May 16 13:15:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes
May 16 13:14:54 vespa nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes
May 16 13:09:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;OK;SOFT;2;PROCS OK: 250 processes
May 16 13:08:54 ank-morpork nagios3: SERVICE ALERT: localhost;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes
May 16 09:18:54 alien nagios3: SERVICE ALERT: localhost;Total Processes;OK;HARD;4;PROCS OK: 241 processes
May 16 09:09:54 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;OK;SOFT;3;OK - load average: 3.37, 2.53, 1.25
May 16 09:08:54 ank-morpork nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;2;WARNING - load average: 7.56, 2.83, 1.25
May 16 09:01:25 alien nagios3: SERVICE ALERT: localhost;Current Load;WARNING;SOFT;1;WARNING - load average: 8.48, 2.65, 1.16
May 19 20:41:11 ank-morpork nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;HARD;4;DISK CRITICAL - free space: /media/cdrom0 0 MB (0% inode=0%):
May 19 20:40:11 cats nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;3;DISK CRITICAL - free space: /media/cdrom0 0 MB (0% inode=0%):
May 19 20:39:11 ank-morpork nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;2;DISK CRITICAL - free space: /media/cdrom0 0 MB (0% inode=0%):
May 19 20:38:11 vespa nagios3: SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;1;DISK CRITICAL - free space: /media/cdrom0 0 MB (0% inode=0%):